# Assignment 0: Practice

This assignment exists for practice only.

This assignment is out of 0 points, worth 0% of your grade.

This assignment covers assignment completion, datahub submission, and anaconda downloading.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.**


## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with 'assert' statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

### Q1 - Practice Question (0 points)

Consider your response to the following question: which course are you MOST excited about this quarter?:
- A) COGS 18
- B) Introduction to Python
- C) COGS18
- D) Intro to Python

Replace `raise NotImplementedError()` below to assign the letter corresponding to your response to the variable `favorite_class`. Be sure that you put quotes around the the letter for your response. 

For example, if "COGS 18" (option A above) were my answer, I would replace `raise NotImplementedError()` with `favorite_class = 'A'`.

After you add your answer, run the code in that cell using 'Shift + Enter' or clicking the "Run" icon above. A number will appear to the left of the cell once your code has executed.

<!-- grade_id: cell-e4f640daee64fc2e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Typically, after you write some code in your assignments, there will be a few assert statements, similar to what you see below. These help to check that you're on the right path. If you are on the right track, when you run the code in those cells, they will "pass silently", meaning you won't see any output. 

However, if you get an AssertionError, you'll want to return to the code you wrote and re-think your response.

<!-- grade_id: cell-629d1512b290862d, points: 0, grade: true, solution: false -->
```python
possible_answers = ['A', 'B', 'C', 'D']
assert favorite_class in possible_answers
```

**Note**: Just because an assert doesn't fail does *not* mean you've got the answer right. There are hidden tests in your assignments. These check to ensure the code you've written does what we asked. As long as you do what we ask *and* pass the assert statements, you'll receive full credit.

## The End!

This is the end of the practice assignment!

Have a look back over your answers, and also make sure to `Restart & Run All` from the kernel menu to double check that everything is working properly. This restarts everything and runs your code from top to bottom.

When you are ready to submit your assignment, you can click validate at the top. Note that in later assignments the code will take too long to run and validation may fail. Validation is just a final check that all the asserts are passing without failing.

Once you're happy with your work, click the disk icon to save, return to the Assignments tab, and submit on DataHub! **You MUST submit the notebook to receive assignment credit. We cannot grade assignments that have not been submitted.**

Note that you can submit at any time, but **we grade your most recent submission**. This means that **if you submit an updated notebook after the submission deadline, it will be marked as late**.